---
title       : Exploring Univariate Data
description : (UNDER CONSTRUCTION) This set of exercises introduces some basic plotting functions in R that can be used to visually inspect the distribution of a single variable.

--- type:NormalExercise lang:r xp:50 skills:1

## Executing Multiple Lines

Up until now we have executed the code in our script one line at a time. It is however possible to select multiple lines and execute them all. This can be done by "painting" the lines and then pressing `Ctrl + Enter` as usual. Painting is done by left clicking your mouse, holding down and then moving the mouse.

Usually in R it does not matter if you execute the lines of a script together or separately from top to bottom, but here in datacamp it sometimes does. For example drawing a plot and adding information to it has to be done in a single execution.

*** =instructions
- Execute the code in the script normally one line of a time. What happens?
- Select both lines by painting them with your mouse and execute the code again.


*** =pre_exercise_code
```{r}
# pre exercise code here
```


*** =sample_code
```{r}
# draw a plot and then add a horizontal line
plot(x=1, y=2)
abline(h=2)

```


*** =sct
```{r}

# test error
test_error()


# Final message the student will see upon completing the exercise
success_msg("Good work!")

```


--- type:NormalExercise lang:r xp:100 skills:1

## USA state data

In this set of exercises we will be using data that is stored in an R object `US`, which is of the familiar class data frame. This data contains information about the states of the USA.  

The data is taken from the 'state' data set that comes with R. As with functions, the data sets in R have help pages that you can access with `?` or `help()`. 
  
  *** =instructions
- Check out the structure of the `US` object
- Write `?state` or `help(state)` to the console to open up the help page on `state`. Take a look at the variable descriptions under `state.x77`. This is where all but one variable of `US` is described.
- Create a summary of the variables in `US`

*** =hint
- `str()`-function gives information about R-objects
- `summary()` creates summaries of the columns of a data.frame.
  
  *** =pre_exercise_code
```{r}
# create a data frame object US
US <- data.frame(State = row.names(state.x77))
US <- cbind(US, as.data.frame(state.x77, row.names=F))

# add a population density (a factor)
pop_area <- with(US, Population/Area)
breaks <- quantile(pop_area, probs=c(0,0.33,0.66,1))
US$PoptoArea <- factor(cut(pop_area, breaks=breaks, include.lowest = T),
                      labels=c("low","medium","high"))

colnames(US)[c(5,7)] <- c("Life_Exp","HS_Grad")
```

*** =sample_code
```{r}
# US data frame is available


# Check out the structure of US.



# Create a summary of the variables in US.







# The console is below
```

*** =solution
```{r}

# Check out the structure of US.
str(US)


# Create a summary of the variables in US.
summary(US)


```

*** =sct
```{r}

test_function("str", args = "object")
test_function("summary", args = "object" )

# test for error
test_error()

success_msg("Great work! Now let's do some plotting :)")

```

--- type:NormalExercise lang:r xp:150 skills:1

## Bar plots

Bar plots are useful for visualizing grouped data. A simple example is the level counts of a factor.  

In base R graphics, the function `barplot()` can be used to draw a bar plot. It's first argument is `height` which is the only mandatory argument. The information for `height` can be computed easily by creating a `summary()` or a `table()` of a factor.

In this exercise you will draw a barplot of `PoptoArea` in the data frame `US`. `PoptoArea` describes the population to area ratio in a state. You can use `str()` to see the structure of `PoptoArea`.

*** =instructions
- Remind yourself on how to access the column of a data frame and how to create objects.
- Use `table()` or `summary()` on `PoptoArea` in the data frame `US`. What do they show you?
- Create an object `tmp` by assigning the output of `table()` or `summary()` to it.
- Use the function `barplot()` to draw a bar plot. Use the object you created as the first argument.

*** =hint
- You can access the columns of a data.frame (df) with the `$` operator: `df$column`.
- Remember how to create objects using the assign operator `<-`: `object_name <- function_call()`


*** =pre_exercise_code
```{r}
# create a data frame object US
US <- data.frame(State = row.names(state.x77))
US <- cbind(US, as.data.frame(state.x77, row.names=F))

# add a population density (a factor)
pop_area <- with(US, Population/Area)
breaks <- quantile(pop_area, probs=c(0,0.33,0.66,1))
US$PoptoArea <- factor(cut(pop_area, breaks=breaks, include.lowest = T),
                      labels=c("low","medium","high"))

colnames(US)[c(5,7)] <- c("Life_Exp","HS_Grad")
```


*** =sample_code
```{r}
# US data frame is available


# Use table() or summary() on PoptoArea in US.


# Create object tmp that stores the information.



# Call barplot() to draw a graph. Use the object you created as an argument.




```

*** =solution
```{r}
# Look at the counts of different levels of Pop_Area in US.
table(US$PoptoArea)


# Create object tmp that stores the counts.
tmp <- table(US$PoptoArea)


# Call barplot() to draw a graph. Use the object you created as an argument.
barplot(tmp)



```

*** =sct
```{r}

test_function("barplot", args="height")

# test error
test_error()

# Final message the student will see upon completing the exercise
success_msg("Excellent work!")

```

--- type:NormalExercise lang:r xp:100 skills:1

## Adding labels to plots

Our bar plot looks good but it could use some extra information. Often you want to add a main title to the plot and perhaps add a label to the y or x axis.

You can easily add to a plot by calling further functions. For example you can use the `title()` function to add labels defined as character strings.  

Here in datacamp the functions that add to the plot need to be called simultaneously with the ones that draw the plot.

*** =instructions
- Open the help page of `title()`.
- Figure out which parameter of `title()` corresponds to the main title.
- Add a call to title `title()` and give you graph the main title: "My barplot". Remember to use quotation marks.

*** =hint
- `?title` or `help(title)` opens a help page.
- The first argument, `main`, of `title()` can be used to set the main title.


*** =pre_exercise_code
```{r}
# create a data frame object US
US <- data.frame(State = row.names(state.x77))
US <- cbind(US, as.data.frame(state.x77, row.names=F))

# add a population density (a factor)
pop_area <- with(US, Population/Area)
breaks <- quantile(pop_area, probs=c(0,0.33,0.66,1))
US$PoptoArea <- factor(cut(pop_area, breaks=breaks, include.lowest = T),
                      labels=c("low","medium","high"))

colnames(US)[c(5,7)] <- c("Life_Exp","HS_Grad")
barplot(table(US$PoptoArea))
```


*** =sample_code
```{r}
# US data frame is available

# The function calls below draw a barplot and add an x-axis label.
barplot(table(US$PoptoArea))
title(xlab="population/area")

# Add a call to title() that adds a main title to the plot. 



```

*** =solution
```{r}

# The function calls below draw a barplot and add an x-axis label.
barplot(table(US$PoptoArea))
title(xlab="population/area")

# Add a call to title() that adds a main title to the plot.
title("My barplot")


```

*** =sct
```{r}

# test error
test_error()


# Final message the student will see upon completing the exercise
success_msg("Great job!")

```

--- type:NormalExercise lang:r xp:100 skills:1

## Advanced summary statistics: variance

We have already seen the usefulness of the `summary()` function. However there are many more summary statistics that might interest us that are not reported in the default summary. Good examples are the variance and the standard deviation. 

The variance can be computed with the `var()`-function. The first argument `x` should be a numeric vector, such as a column of a data frame.  

The argument `x` of `var()` can also be a numeric matrix or a data frame to compute <a href="https://en.wikipedia.org/wiki/Covariance">covariances</a>. For this there is also the function `cov()`.

*** =instructions
- Create a summary of the `Income` column in data frame `US`.
- Compute the variance of `Income`. Save the result to an object called `inc_var`
- Use any resource to find the connection between variance and standard deviation.
- Use `inc_var` to compute the standard deviation of `Income`. Save the result to an object called `inc_sd`.
- Print the rounded values of both objects. How would you interpret the measures?
- Do you think you could compute the standard deviation with a single function call?

*** =hint
- You can access a column of a data.frame (df) by `df$column_name`.
- You can use the assign-operator `<-` to create an object.
- Standard deviation is the square root of the variance.
- You can type and execute the objects name to print it's value.


*** =pre_exercise_code
```{r}
# create a data frame object US
US <- data.frame(State = row.names(state.x77))
US <- cbind(US, as.data.frame(state.x77, row.names=F))

# add a population density (a factor)
pop_area <- with(US, Population/Area)
breaks <- quantile(pop_area, probs=c(0,0.33,0.66,1))
US$PoptoArea <- factor(cut(pop_area, breaks=breaks, include.lowest = T),
                      labels=c("low","medium","high"))

colnames(US)[c(5,7)] <- c("Life_Exp","HS_Grad")
```


*** =sample_code
```{r}
# US data frame is available

# The function sqrt() computes the squre root and the function round() rounds to specified digits (default is digits = 0)
ten_squared <- sqrt(10)
ten_squared # = 3.162...
round(ten_squared) # = 3

# Create a summary of Income in US.


# Create object inc_var. Use var() to compute the variance.


# Create object inc_sd.


# Print the rounded values of both objects.



```

*** =solution
```{r}
# Create a summary of Income in US.
summary(US$Income)

# Create object inc_var.
inc_var <- var(US$Income)

# Create object inc_sd
inc_sd <- sqrt(inc_var)

# Print the rounded values of both objects.
round(inc_var)
round(inc_sd)

```

*** =sct
```{r}

test_function("summary",args="object")
test_object("le_var")
test_object("le_sd")

test_output_contains(round(inc_var))
test_output_contains(round(inc_sd))

# test error
test_error()

# Final message the student will see upon completing the exercise
success_msg("Great work!")

```


--- type:NormalExercise lang:r xp:100 skills:1

## Histograms

Histograms are an excellent way of visualizing the shape of a distribution. In base R, histograms are drawn with the `hist()` function. The first argument of `hist()`, `x` should be a numeric vector of data values. The columns of a data frame can be such vectors, for example.

As with all the plotting functions in R, the `hist()` function can take on many additional arguments which allow for more control over the plot. One simple argument is `col`, which defines the color. `col` can be a character string giving the name of the color ("blue", "green", ...). Some possible color names can be found <a target="_blank" href="http://www.stat.columbia.edu/~tzheng/files/Rcolor.pdf">here</a>

*** =instructions
- Create a histogram of the Incomes in the USA states.
- Adjust your code and give your histogram a nice color using the `col` argument. 


*** =hint
- You can use `str(US)` to see the structure (and columns) of `US`
- You can access the column of a data frame (df) with the `$` symbol: `df$column_name`.


*** =pre_exercise_code
```{r}
# create a data frame object US
US <- data.frame(State = row.names(state.x77))
US <- cbind(US, as.data.frame(state.x77, row.names=F))

# add a population density (a factor)
pop_area <- with(US, Population/Area)
breaks <- quantile(pop_area, probs=c(0,0.33,0.66,1))
US$PoptoArea <- factor(cut(pop_area, breaks=breaks, include.lowest = T),
                      labels=c("low","medium","high"))

colnames(US)[c(5,7)] <- c("Life_Exp","HS_Grad")
```


*** =sample_code
```{r}
# US data frame is available

# Draw a histogram of Income in US. Make it a colored one!






```

*** =solution
```{r}
# Draw a histogram of Income in US. Make it a colored one!
hist(US$Income, col ="blue")





```

*** =sct
```{r}

test_function("hist", args="x")

# test error
test_error()


# Final message the student will see upon completing the exercise
success_msg("Good work!")

```

--- type:NormalExercise lang:r xp:150 skills:1

## Box plots

One very good way to summarize and compare distributions is by drawing box plots. A box plot visualizes the 25th, 50th and 75th percentiles (the box), the typical range (the whiskers) and the outliers (the dots) of a distribution. 

The whiskers extending from the box can be computed by several techniques. The default in R is to extend them to reach to a data point that is no more than 1.5*IQR away from the box, where IQR is the inter quartile range defined as  

`IQR = 75th percentile - 25th percentile`  

Values outside the whiskers can be considered as outliers, unusually distant observations. For more information on IQR, see <a target="_blank" href ="https://en.wikipedia.org/wiki/Interquartile_range"> wikipedia</a>

*** =instructions
- Draw a `boxplot()` of `US$Income`. The first argument (`x`) of `boxplot()` should be a numerical vector containing data values.
- Label the y axis, for example: "income per capita"
- Add a main title, for example: "Distribution of per capita income in USA states".

*** =hint
- Give the `Income` column of `US` as the first argument to `boxplot` to draw the plot.
- Use `title()` to add labels defined as characted strings. See ?title for help about argument names etc.

*** =pre_exercise_code
```{r}
# create a data frame object US
US <- data.frame(State = row.names(state.x77))
US <- cbind(US, as.data.frame(state.x77, row.names=F))

# add a population density (a factor)
pop_area <- with(US, Population/Area)
breaks <- quantile(pop_area, probs=c(0,0.33,0.66,1))
US$PoptoArea <- factor(cut(pop_area, breaks=breaks, include.lowest = T),
                      labels=c("low","medium","high"))

colnames(US)[c(5,7)] <- c("Life_Exp","HS_Grad")
```


*** =sample_code
```{r}
# US data frame is available

# Draw a boxplot of Income in US. Add a y axis label and a main title.





```

*** =solution
```{r}
# Draw a boxplot of Income in US. Add a y axis label and a main title.
boxplot(US$Income)
title(ylab="income per capita")
title("Distribution of per capita income in USA states")


```

*** =sct
```{r}


test_function("boxplot", args="x")

# test error
test_error()


# Final message the student will see upon completing the exercise
success_msg("Awsome work!")

```
